---
title: PCA结果展示
description: 这是一个测试的章节
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise. This is an example exercise.

`@instructions`
This is an example exercise. This is an example exercise.

`@hint`
This is an example exercise. This is an example exercise.

`@pre_exercise_code`
```{r}
install.package('tidyverse')
library(tidyverse)
```

`@sample_code`
```{r}
print(1)
```

`@solution`
```{r}
print(1)
```

`@sct`
```{r}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: b9221109fe
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
